```python
from nose.tools import assert_equal, assert_true, assert_false
```

Use list comprehension to create a list of the first five square integers (starting from 1).

```Python
[1, 4, 9, 16, 25]
```

<!-- grade_id: cell-c7e1f10b26c3abf2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Use list comprehension to create a list of two-tuples describing the integer coordinates $(x,y)$ of a 2D grid $-3 \le y \le 3, -2 \le x \le 4$. 
```Python
[(-2,-3), (-2,-2), ..., (4,2), (4,3)]
```

<!-- grade_id: cell-c8c2f16a71f06f66, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Use list comprehension to create a list of two-tuples describing the integer coordinates $(x,y)$ of a 2D grid  where $ -2 \le y \lt x, -2 \le x \le 4$. 
```Python
[(-1,-2), (0,-2), (0,-1),..., (4,3)]
```

<!-- grade_id: cell-c0f3eb34dd482642, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Write a function `merge` that takes a list of lists (e.g. `[[1,3],['one','three'],[(4,5),(6,7)]]` and returns a list with all the values in a single list (e.g. `[[1,3],['one','three'],[(4,5),(6,7)]]`->`[1,3,'one','three',(4,5), (6,7)]`

<!-- grade_id: cell-ad562715127a5e97, points: 0, grade: false, solution: true -->
```python
def merge(mylist):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f0fce91efed04932, points: 0, grade: true, solution: false -->
```python
assert_equal(merge([[1,2,3],[4,5,6],[(1,2), (3,4), (5,6)]]), [1, 2, 3, 4, 5, 6, (1, 2), (3, 4), (5, 6)])
```

The *proper factors* for an integer $N$ are those integers $f$ that divide evenly into $N$ (i.e. with zero remainder) such that $1 \le f \lt N$. Write a function `pfactors` that takes as an argument an integer `N` and returns a list with the proper factors of `N`.

<!-- grade_id: cell-530ea9610c5f506d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-12b150062c1a9814, points: 0, grade: true, solution: false -->
```python
assert_equal(pfactors(9),[1,3])
assert_equal(pfactors(6),[1,2,3])
assert_equal(pfactors(12), [1,2,3,4,6])
```

An integer $N$ is called a perfect number of the sum of its proper factors equals $N$. Write a function `perfect` that takes an integer argument and returns `True` if $N$ is perfect and `False` otherwise.

<!-- grade_id: cell-f732571db171e67e, points: 0, grade: false, solution: true -->
```python
def perfect(N):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7d2c6ad4ca22e628, points: 0, grade: true, solution: false -->
```python
assert_true(perfect(28))
assert_false(perfect(9))
assert_true(perfect(6))
assert_false(perfect(12))

```

Using the function `perfect` and list comprehension, write a function `perfects` that returns all the perfect numbers less than some integer `N`.

<!-- grade_id: cell-df3fc8e5b03306b1, points: 0, grade: false, solution: true -->
```python
def perfects(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
assert_equal(perfects(10), [6])
assert_equal(perfects(100), [6, 28])
assert_equal(perfects(1000), [6, 28, 496])
assert_equal(perfects(10000), [6, 28, 496, 8128])
```

Using a string `split()` method, the `zip` function and list comprehension, write a function `bigrams` that returns as a list of tuples all the `bi-grams` (two consecutive words, for example ("we", "present")). Text is taken from [this](https://www.ncbi.nlm.nih.gov/pubmed/29610178) Pubmed abstract.

```python
txt = """We present a case of an infant with proximal hypospadias, penoscrotal transposition, and bilaterally 
descended testes found to have a clinically significant WT1 gene alteration on a customized disorder of sex 
development genetic panel in which 62 genes associated with 46, XY disorders of sex development were evaluated. 
This diagnosis led to early screening for and diagnosis and treatment of Wilms tumor. Patients with proximal 
hypospadias are not routinely evaluated by genetic testing, and when initial hormonal analyses are within normal 
ranges for a typical male patient, the genital atypia is usually attributed to an isolated anatomic abnormality. 
There is no consensus among urologists, endocrinologists, or geneticists regarding when genetic testing is warranted 
in these patients or the extent of genetic testing that should be pursued. However, given advances in genetic testing 
and the discovery of more genetic variants, the genetic evaluation of infants with proximal hypospadias should be 
considered on an individual patient basis. Only with continued evaluation and the identification of further genetic 
variants can we establish future parameters for genetic evaluation in patients with proximal hypospadias 
and more appropriately counsel patients and their families regarding the implications of these variants."""
```

<!-- grade_id: cell-94f2878993abb3da, points: 0, grade: false, solution: true -->
```python
def bigrams(txt):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7e1e03b761cbc202, points: 0, grade: true, solution: false -->
```python
assert_equal(bigrams(txt)[15], ('testes', 'found'))

assert_equal(bigrams(txt)[38], ('associated', 'with'))
```

Use list comprehension to write a function `gcount` to count how many guanine nucleotides are in the sequence (Carpodacus mexicanus) provided below.

<!-- grade_id: cell-00ccffc957bcd080, points: 0, grade: false, solution: true -->
```python
seq = """CCTTTATCTAATCTTTGGAGCATGAGCTGGCATAGTTGGAACCGCCCTCAGCCTCCTCATCCGTGCAGAA
CTTGGACAACCTGGAACTCTTCTAGGAGACGACCAAATTTACAATGTAATCGTCACTGCCCACGCCTTCG
TAATAATTTTCTTTATAGTAATACCAATCATGATCGGTGGTTTCGGAAACTGACTAGTCCCACTCATAAT
CGGCGCCCCCGACATAGCATTCCCCCGTATAAACAACATAAGCTTCTGACTACTTCCCCCATCATTTCTT
TTACTTCTAGCATCCTCCACAGTAGAAGCTGGAGCAGGAACAGGGTGAACAGTATATCCCCCTCTCGCTG
GTAACCTAGCCCATGCCGGTGCTTCAGTAGACCTAGCCATCTTCTCCCTCCACTTAGCAGGTGTTTCCTC
TATCCTAGGTGCTATTAACTTTATTACAACCGCCATCAACATAAAACCCCCAACCCTCTCCCAATACCAA
ACCCCCCTATTCGTATGATCAGTCCTTATTACCGCCGTCCTTCTCCTACTCTCTCTCCCAGTCCTCGCTG
CTGGCATTACTATACTACTAACAGACCGAAACCTAAACACTACGTTCTTTGACCCAGCTGGAGGAGGAGA
CCCAGTCCTGTACCAACACCTCTTCTGATTCTTCGGCCATCCAGAAGTCTATATCCTCATTTTAC"""
def gcount(seq):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c09da04a2dd9fb40, points: 0, grade: true, solution: false -->
```python
assert_equal(gcount(seq), 108)
```

```python

```