```python
import random
import seaborn as sns
from nose.tools import assert_equal, assert_true, assert_false, assert_almost_equal
```

```python
class die(object):
    def __init__(self, n=6):
        self.__values = list(range(1,n+1))
    def roll(self):
        return random.choice(self.__values)
    
    @property
    def values(self):
        return self.__values[:]
    @values.setter
    def values(self, values):
        self.__values = values
```

Write a class `unfairDie` that inherits from `die` but represents anunfair die.

**Hints**

In my solution I used the `random.choices` and `random.randint` (with arguments 1 and 20) functions. I also used list concatenation.

<!-- grade_id: cell-99902cf4d21eef86, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
die1 = unfairDie(n=6)
rolls = [die1.roll() for i in range(100000)]
```

```python
sns.distplot(rolls,kde=False, norm_hist=True)
```

Write a function `count` that takes a list of roll values and returns a dictionary with an estimation of the probability of each value being rolled. Keys should be the die values (i.e. 1 to 6) and the values should be the estimated probability of that value being rolled. Compare this results to what you would compute from the values property of the unfair die.

<!-- grade_id: cell-41fe19c68c130180, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
count(rolls)
```

<!-- grade_id: cell-ea6de06c76077dd3, points: 0, grade: true, solution: false -->
```python
assert_almost_equal(sum(count(rolls).values()), 1.0)
```

Use recursion to write a function `double_every_other` that takes as an argument a sequence of numbers (list or tuple) and returns a tuple with ever other element (starting from the second element) doubled (e.g. [0,1,2,3,4,5]->[0,2,2,6,4,10]).

**Hints:** 

There will be two base cases:

1. When the list is empty
2. When the list has only one element in it

Use slicing and concatenation.

<!-- grade_id: cell-0e669bc20815a9f7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-798764d5e1611e8d, points: 0, grade: true, solution: false -->
```python
assert_equal(double_every_other([1,2,3,4]), (1, 4, 3, 8))
assert_equal(double_every_other([7,3,1,5,0,2,-3,4]), (7, 6, 1, 10, 0, 4, -3, 8))
```

The factorial function (usually denoted with a $!$ as in $5!$) is defined as 
\begin{equation}
n! = 1\times 2 \times \cdots \times (n-1) \times n
\end{equation}

Defining $0!=1$, we can express the factorial function recursively

$$
 n! = 
  \begin{cases} 
   0 & \text{if } n=0 \\
   n\cdot(n-1)!       & \text{otherwise }
  \end{cases}
$$

<!-- grade_id: cell-f1f6960566eac179, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3647176a91d5d012, points: 0, grade: true, solution: false -->
```python
assert_equal(factorial(5), 120)
assert_equal(factorial(23),25852016738884976640000)
```

Write a function `product` that computes the product of all the values in a list.

<!-- grade_id: cell-0aa434ef4196b47f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a286bd335ae68929, points: 0, grade: true, solution: false -->
```python
assert_equal(product([1,2,3]),6)
assert_equal(product([1,1,3,4,3,2]),72)
```

Use recursion to write a function `euclid` that comptues the greatest common denominator of two positive integers using Euclids algorithm for the GCD.

<!-- grade_id: cell-d8b2025c42ce7e46, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c44ec6e6f9dd4261, points: 0, grade: true, solution: false -->
```python
assert_equal(euclid(27,6),3)
assert_equal(euclid(12537,35),7)
```

Use recursion to write a Python implementation of the quick sort algorithm to sort the values of a list.

<!-- grade_id: cell-e47f437a781382cb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d10e8fb5e0b4df17, points: 0, grade: true, solution: false -->
```python
assert_equal(qsort([3,1,5,4,-2,9,17,-4]), [-4, -2, 1, 3, 4, 5, 9, 17])
assert_equal(qsort([3,1,1,2,3,5,4,7]), [1, 1, 2, 3, 3, 4, 5, 7])
```

Write a function `drop` that takes a list and and an integer N and returns a list with the first N elements dropped.

<!-- grade_id: cell-7aad3abe169ad963, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b4f14c2b0ac62696, points: 0, grade: true, solution: false -->
```python
assert_equal(drop([1,2,3,4,5],2), [3,4,5])
assert_equal(drop([5,4,3,2,1],0), [5, 4, 3, 2, 1])
assert_equal(drop([1,2,3,4,5],9),[])
```

Write a function `zipper` that takes two sequences, `xs` and `ys`, and returns a sequence of the corresponding pairs ---[(x[0],ys[0]), (xs[1],ys[1]),...]---from `xs` and `ys` until one of the sequences is exhausted.

<!-- grade_id: cell-2f8dd69e7c628c8a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e579ddbca10a7fa1, points: 0, grade: true, solution: false -->
```python
assert_equal(zipper([1,2,3],[1,2]), [(1, 1), (2, 2)])
assert_equal(zipper([1,2],[1,2,3]), [(1, 1), (2, 2)])
assert_equal(zipper([1,2,3],[]),[])
assert_equal(zipper([1,2,3,4],[4,3,2,1]), [(1, 4), (2, 3), (3, 2), (4, 1)])
assert_equal(zipper("Brian", "Chapman"), [('B', 'C'), ('r', 'h'), ('i', 'a'), ('a', 'p'), ('n', 'm')])
```

```python

```